<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml">
  <head>
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>13.8. Programming Exercises &#8212; How to Think like a Computer Scientist: Interactive Edition</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/video.css" type="text/css" />
    <link rel="stylesheet" href="../_static/clickable.css" type="text/css" />
    <link rel="stylesheet" href="../_static/codemirror.css" type="text/css" />
    <link rel="stylesheet" href="../_static/activecode.css" type="text/css" />
    <link rel="stylesheet" href="../_static/datafile.css" type="text/css" />
    <link rel="stylesheet" href="../_static/parsons.css" type="text/css" />
    <link rel="stylesheet" href="../_static/lib/prettify.css" type="text/css" />
    <link rel="stylesheet" href="../_static/tabbedstuff.css" type="text/css" />
    <link rel="stylesheet" href="../_static/dragndrop.css" type="text/css" />
    <link rel="stylesheet" href="../_static/showEval.css" type="text/css" />
    <link rel="stylesheet" href="../_static/fitb.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pytutor.css" type="text/css" />
    <link rel="stylesheet" href="../_static/modal-basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/poll.css" type="text/css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript">
      var DOCUMENTATION_OPTIONS = {
        URL_ROOT:    '../',
        VERSION:     '3.0',
        COLLAPSE_INDEX: false,
        FILE_SUFFIX: '.html',
        HAS_SOURCE:  true,
        SOURCELINK_SUFFIX: '.txt'
      };
    </script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="14. Dictionaries" href="../CMPS5P/l14.html" />
    <link rel="prev" title="13.7. Glossary" href="Glossary.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'thinkcspy';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "4.0.1";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = true;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'thinkcspy';
  eBookConfig.runestone_version = '3.0.5';
</script>
<style>
.navLink{
	display: block;
  position: static;
  font-size: 3rem;
}

@media (max-width: 600px) {
    /* Styling definitions for chevron function in bookfuncs.js */
	.navLink{
		display: block;
	}
}

#relations-next {
  float: right;
}

#relations-prev {
  float: left;
}

.navLink a {
	color: #999;
}
</style>

  </head>
  <body>



<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="programming-exercises">
<h1>13.8. Programming Exercises<a class="headerlink" href="#programming-exercises" title="Permalink to this headline">¶</a></h1>

    <div data-component="question" class="full-width container question" id="recursion_ex_1" >
    <ol start=1 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_1" class="runestone explainer ac_section alert alert-warning">
Write a recursive function to compute the factorial of a number.
<textarea data-component="activecode" id=ex_rec_1 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def computeFactorial(number):
    #your code here

====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

    def testOne(self):
        self.assertEqual(computeFactorial(0),1,"Tested computeFactorial on input 0")
        self.assertEqual(computeFactorial(1),1,"Tested computeFactorail on input 1")
        self.assertEqual(computeFactorial(2),2,"Tested computeFactorial on input 2")
        self.assertEqual(computeFactorial(3),6,"Tested computeFactorial on input 3")
        self.assertEqual(computeFactorial(4),24,"Tested computeFactorial on input 4")
        self.assertEqual(computeFactorial(8),40320,"Tested computeFactorial on input 8")
        self.assertEqual(computeFactorial(-5),None,"Tested computeFactorial on a negative number - make sure to handle this case")

myTests().main()
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_2" >
    <ol start=2 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_2" class="runestone explainer ac_section alert alert-warning">
Write a recursive function to reverse a list.
<textarea data-component="activecode" id=ex_rec_2 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >
def reverseList(lst):
    #your code here


====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):
    def testOne(self):
        self.assertEqual(reverseList([1,2,3,4,5]), [5,4,3,2,1], "Your function failed with input [1,2,3,4,5]")
        self.assertEqual(reverseList(['Hello','World','!']), ['!','World','Hello'], "Your function failed with input ['Hello,'World','!']")
        self.assertEqual(reverseList(['Python',100,'35','Computer Science']), ['Computer Science', '35', 100, 'Python'], "Your function failed with input ['Python,100,'35','Computer Science']")

myTests().main()
</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_3" >
    <ol start=3 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_3" class="runestone explainer ac_section alert alert-warning">
<p>Modify the recursive tree program using one or all of the following
ideas:</p>
<ul class="simple">
<li>Modify the thickness of the branches so that as the <code class="docutils literal"><span class="pre">branchLen</span></code>
gets smaller, the line gets thinner.</li>
<li>Modify the color of the branches so that as the <code class="docutils literal"><span class="pre">branchLen</span></code> gets
very short it is colored like a leaf.</li>
<li>Modify the angle used in turning the turtle so that at each branch
point the angle is selected at random in some range. For example
choose the angle between 15 and 45 degrees. Play around to see
what looks good.</li>
<li>Modify the <code class="docutils literal"><span class="pre">branchLen</span></code> recursively so that instead of always
subtracting the same amount you subtract a random amount in some
range.</li>
</ul>
<p>If you implement all of the above ideas you will have a very
realistic looking tree.</p>

<textarea data-component="activecode" id=ex_rec_3 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_4" >
    <ol start=4 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_4" class="runestone explainer ac_section alert alert-warning">
Find or invent an algorithm for drawing a fractal mountain. Hint: One
approach to this uses triangles again.
<textarea data-component="activecode" id=ex_rec_4 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_5" >
    <ol start=5 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_5" class="runestone explainer ac_section alert alert-warning">
Write a recursive function to compute the Fibonacci sequence. How
does the performance of the recursive function compare to that of an
iterative version?
<textarea data-component="activecode" id=ex_rec_5 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_6" >
    <ol start=6 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_6" class="runestone explainer ac_section alert alert-warning">
Implement a solution to the Tower of Hanoi using three stacks to keep
track of the disks.
<textarea data-component="activecode" id=ex_rec_6 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_7" >
    <ol start=7 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_7" class="runestone explainer ac_section alert alert-warning">
Using the turtle graphics module, write a recursive program to
display a Hilbert curve.
<textarea data-component="activecode" id=ex_rec_7 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_8" >
    <ol start=8 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_8" class="runestone explainer ac_section alert alert-warning">
Using the turtle graphics module, write a recursive program to
display a Koch snowflake.
<textarea data-component="activecode" id=ex_rec_8 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_9" >
    <ol start=9 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_9" class="runestone explainer ac_section alert alert-warning">
Write a program to solve the following problem: You have two jugs: a
4-gallon jug and a 3-gallon jug. Neither of the jugs have markings on
them. There is a pump that can be used to fill the jugs with water.
How can you get exactly two gallons of water in the 4-gallon jug?
<textarea data-component="activecode" id=ex_rec_9 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_10" >
    <ol start=10 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_10" class="runestone explainer ac_section alert alert-warning">
Generalize the problem above so that the parameters to your solution
include the sizes of each jug and the final amount of water to be
left in the larger jug.
<textarea data-component="activecode" id=ex_rec_10 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_11" >
    <ol start=11 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_11" class="runestone explainer ac_section alert alert-warning">
Write a program that solves the following problem: Three missionaries
and three cannibals come to a river and find a boat that holds two
people. Everyone must get across the river to continue on the
journey. However, if the cannibals ever outnumber the missionaries on
either bank, the missionaries will be eaten. Find a series of
crossings that will get everyone safely to the other side of the
river.
<textarea data-component="activecode" id=ex_rec_11 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_12" >
    <ol start=12 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_12" class="runestone explainer ac_section alert alert-warning">
Modify the Tower of Hanoi program using turtle graphics to animate
the movement of the disks. Hint: You can make multiple turtles and
have them shaped like rectangles.
<textarea data-component="activecode" id=ex_rec_12 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true"  data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    
    <div data-component="question" class="full-width container question" id="recursion_ex_13" >
    <ol start=13 class=arabic><li class="alert alert-warning">

    
<div data-childcomponent="ex_rec_13" class="runestone explainer ac_section alert alert-warning">
<p>Pascal’s triangle is a number triangle with numbers arranged in
staggered rows such that</p>
<div class="math">
\[a_{nr} = {n! \over{r! (n-r)!}}\]</div>
<p>This equation is the equation for a binomial coefficient. You can
build Pascal’s triangle by adding the two numbers that are diagonally
above a number in the triangle. An example of Pascal’s triangle is
shown below.</p>
<div class="highlight-default"><div class="highlight"><pre><span></span>        <span class="mi">1</span>
      <span class="mi">1</span>   <span class="mi">1</span>
    <span class="mi">1</span>   <span class="mi">2</span>   <span class="mi">1</span>
  <span class="mi">1</span>   <span class="mi">3</span>   <span class="mi">3</span>   <span class="mi">1</span>
<span class="mi">1</span>   <span class="mi">4</span>   <span class="mi">6</span>   <span class="mi">4</span>   <span class="mi">1</span>
</pre></div>
</div>
<p>Write a program that prints out Pascal’s triangle. Your program
should accept a parameter that tells how many rows of the triangle to
print.</p>

<textarea data-component="activecode" id=ex_rec_13 data-lang="python"  data-hidecode="true"  data-timelimit=25000 data-coach="true" data-codelens="true" data-audio=''    data-gradebutton=true >

</textarea>
</div>

    </li></ol>
    </div>
    </div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 13.7. Glossary' data-toggle="tooltip">
    <a href="Glossary.html" >
      <i class='prevNav glyphicon glyphicon-chevron-left'></i>
    </a>
  </li>

  <li class="navLink" style="text-align: center" title='13. Recursion' data-toggle="tooltip" >
    <a href="../CMPS5P/l13.html">
      <i id="relationsNextIcon"  class='nextNav glyphicon glyphicon-chevron-up'></i>
    </a>
  </li>


<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<p style="line-height: 2em">&nbsp;</p>

  </body>
</html>